---
title: 'Programming Basics'
description: 'Chapter description goes here.'
free_preview: true
---

## Welcome to Python Fundamentals

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

In this course we will be teaching you the fundamentals of python programming as well as data science. This course starts you off at the very beginning and assumes no technical knowledge. The main focus in this course is to teach you python programming at a basic level and then we will continue to build on those skills in order to turn you into a complete data scientist where you will know programming, data visualization, linear algebra, statistics, databases, machine learning, and more...

`@instructions`
For now click next to move on to the next lesson.

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
