---
title: 'Programming Basics'
description: 'Chapter description goes here.'
free_preview: true
---

## Welcome to Python Fundamentals For Data Science

```yaml
type: PureMultipleChoiceExercise
key: 79b93128a9
xp: 50
```

In this course we will be teaching you the fundamentals of python programming as well as data science. This course starts you off at the very beginning and assumes no technical knowledge. The main focus in this course is to teach you python programming at a basic level and then we will continue to build on those skills in order to turn you into a complete data scientist where you will know programming, data visualization, linear algebra, statistics, databases, machine learning, and more... Are you excited to begin your journey?

`@hint`


`@possible_answers`
- [Yes]
- No

`@feedback`


---

## Insert exercise title here

```yaml
type: NormalExercise
key: ea7a2a5e06
xp: 100
```

There's been a lot of enthusiasm around data science recently. Using data science, people have been able to do some amazing things with computer programming:

1. AlphaGo through Reinforcement learning to beat the world champion at Go
2. Self-Driving Cars through Deep learning for image recognition
3. Smart Devices which use Natural Language Processing and much more

All this is done by making a computer compute billions of computations. In order to do this we must give the computer instructions. When we do this, we say that we are programming the computer. The instructions we give to the computer must be in a certain programming language and in our case we will be using a programming language called python.

We begin by telling the computer to add two numbers together.

```
12 + 8
```

`@instructions`
Instruct the computer to calculate 12 + 8 in the code editor and then hit submit answer.

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}
12 + 8
```

`@sct`
```{python}

```
